# CS211: Data Privacy
## Homework 4

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

```python
adult.head()
```

Output:
```
Name        DOB          SSN    Zip  Age         Workclass  \
0  Karrie Trusslove   9/7/1967  732-14-6110  64152   39         State-gov   
1  Brandise Tripony   6/7/1988  150-19-2766  61523   50  Self-emp-not-inc   
2     Brenn McNeely   8/6/1991  725-59-9860  95668   38           Private   
3       Dorry Poter   4/6/2009  659-57-4974  25503   53           Private   
4       Dick Honnan  9/16/1951  220-93-3811  75387   28           Private   

   fnlwgt  Education  Education-Num      Marital Status         Occupation  \
0   77516  Bachelors             13       Never-married       Adm-clerical   
1   83311  Bachelors             13  Married-civ-spouse    Exec-managerial   
2  215646    HS-grad              9            Divorced  Handlers-cleaners   
3  234721       11th              7  Married-civ-spouse  Handlers-cleaners   
4  338409  Bachelors             13  Married-civ-spouse     Prof-specialty   

    Relationship   Race     Sex  Capital Gain  Capital Loss  Hours per week  \
0  N
```

## Question 1 (10 points)

Complete the definition of `dp_sum_capgain` below. Your definition should compute a differentially private sum of the "Capital Gain" column of the `adult` dataset, and have a total privacy cost of `epsilon`.

<!-- grade_id: cell-6a369f61d8eef943, points: 0, grade: false, solution: true -->
```python
def dp_sum_capgain(epsilon):        
    # Calclate DP sum
    b = 99999 # essentially the upper bound of our data
    clipped_sum = adult['Capital Gain'].clip(lower=0, upper=b).sum()
    noisy_sum = laplace_mech(clipped_sum, b+1, epsilon)
    
    return noisy_sum

print(adult['Capital Gain'].sum(), dp_sum_capgain(1)) # Keeping epsilon at .1, could be adjusted 



# TRYING TO FIGURE OUT WHAT THE PRIVACY PARAMETER SHOULD BE
bs = range(min(adult['Capital Gain']), max(adult['Capital Gain']), 5) # For performance purposes count by 100
real_sum = adult['Capital Gain'].sum()

for b in bs:
    clipped_sum = adult['Capital Gain'].clip(lower=0, upper=b).sum()
    noisy_sum = laplace_mech(clipped_sum, b, 0.1)
    print('b:', b, 'sum:', pct_error(real_sum, noisy_sum))

# Seems like the percent error continues to drop the closer we get to the max value
# So to be safe that will be our sensitivity!
# Total privacy cost will equal 100,000 * epsilon since we are running 100,000 sum queries

'''
This paragraph has me thinking that we shouldn't even be clipping this summation query at all since 
there is a clear upper and lower bound.

As a rule of thumb, summation queries have unbounded sensitivity when no lower and upper bounds exist on the value of the 
attribute being summed. When lower and upper bounds do exist, the sensitivity of a summation query is equal to the difference 
between them. In the next section, we will see a technique called clipping for enforcing bounds when none exist, 
so that summation queries with unbounded sensitivity can be converted into queries with bounded sensitivity.
'''
```

Output:
```
35089324 35098030.43568285
b: 0 sum: 100.0
b: 5 sum: 99.96151915520711
b: 10 sum: 99.9230252710292
b: 15 sum: 99.88415671575777
b: 20 sum: 99.84558907403132
b: 25 sum: 99.8043355750541
b: 30 sum: 99.76750413158311
b: 35 sum: 99.72920114729588
b: 40 sum: 99.69109392021313
b: 45 sum: 99.64929032131269
b: 50 sum: 99.6192734210969
b: 55 sum: 99.57418537538565
b: 60 sum: 99.5384643636964
b: 65 sum: 99.49806742708044
b: 70 sum: 99.45885728919687
b: 75 sum: 99.4231668513038
b: 80 sum: 99.37624285218226
b: 85 sum: 99.3406305007056
b: 90 sum: 99.30371464472452
b: 95 sum: 99.27019734514637
b: 100 sum: 99.2284851595661
b: 105 sum: 99.19174888954143
b: 110 sum: 99.15335216350338
b: 115 sum: 99.11162024821087
b: 120 sum: 99.07386919956225
b: 125 sum: 99.01237781766697
b: 130 sum: 98.98883463478865
b: 135 sum: 98.9410588015861
b: 140 sum: 98.92166427136387
b: 145 sum: 98.87924694151482
b: 150 sum: 98.84158349199322
b: 155 sum: 98.80293426644755
b: 160 sum: 98.7625130574445
b: 165 sum: 98.72766021710
```

<!-- grade_id: cell-e5b1229ba2249f71, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 1

real_sum = adult['Capital Gain'].sum()
r1 = np.mean([pct_error(real_sum, dp_sum_capgain(0.1)) for _ in range(100)])
r2 = np.mean([pct_error(real_sum, dp_sum_capgain(1.0)) for _ in range(100)])
r3 = np.mean([pct_error(real_sum, dp_sum_capgain(10.0)) for _ in range(100)])

print("Average errors:", r1, r2, r3)

assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 10
assert r2 < 2
assert r3 < 0.2
```

Output:
```
Average errors: 2.9231841477563343 0.27016743182977015 0.027251590405203245
```

## Question 2 (10 points)

In 2-5 sentences each, answer the following:

- What clipping parameter did you use in your definition of `dp_sum_capital`, and why?
- What was the sensitivity of the query you used in `dp_sum_capital`, and how is it bounded?
- Argue that your definition of `dp_sum_capital` has a total privacy cost of `epsilon`

<!-- grade_id: cell-986eec754f32c072, points: 10, grade: true, solution: true -->
- The clipping parameter I chose for the dp_sum_capital function was the max value for the dataset, 99,999. I chose this value because there is no person exhibting more than a value of 99,999 for the capital gain column and thus it is safe to assume 100% of the database is included within the clipping.
- The sensitivity of the query used to compute dp_sum_capital is 100,000. This is because a summation query is unbounded in nature, meaning if a new row were added to the dataset the summation query would then be offset by the amount of the added row. But since our data ranges from 0 to 99,999 it is safe to safe that 100,00 will account for any value in our query that should be added into our dataset, so we can assume a sensitivity of the difference of the lower and upper bounds (with a little padding).
- My definition of dp_sum_capital has a total privacy cost of epsilon. This is because we are only creating one noisy sum here, so the amount of privacy we are using is exactly what is passed in as an argument. Were I to creat a plot or something of the like with the sum function, the total privacy would fluctuate

## Question 3 (10 points)

Complete the definition of `dp_avg_capgain` below. Your definition should compute a differentially private average (mean) of the "Capital Gain" column of the adult dataset, and have a **total privacy cost of epsilon**.

<!-- grade_id: cell-93f2a5153dca0e95, points: 0, grade: false, solution: true -->
```python
def dp_avg_capgain(epsilon):
    real_mean = adult['Capital Gain'].sum() / adult['Capital Gain'].shape[0]
    noisy_mean = dp_sum_capgain(epsilon) / adult['Capital Gain'].shape[0] # just for comparison
    noisier_mean = laplace_mech(dp_sum_capgain(epsilon) / adult['Capital Gain'].shape[0], 1, epsilon)
    
    # sensitivity testing
#     bs = range(0, 100000, 100)
#     real_av = adult['Capital Gain'].mean()

#     for b in bs:
#         clipped_sum = adult['Capital Gain'].clip(lower=0, upper=b).sum()
#         noisy_sum = laplace_mech(clipped_sum, b, 0.1)
#         print('b:', b, 'av:', pct_error(real_sum, noisy_sum))
    
#     print(real_mean, noisy_mean, noisier_mean)
#     d_mean = laplace_mech(adult['Capital Gain'].clip(lower=0, upper=100000).sum(), 99999, epsilon) / laplace_mech(adult['Capital Gain'].clip(lower=0, upper=100000).shape[0], 100000, epsilon)
    return noisier_mean

dp_avg_capgain(1.0)
```

Output:
```
1077.8212240994606
```

<!-- grade_id: cell-ef77b325b5c58908, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 3

real_avg = adult['Capital Gain'].mean()
r1 = np.mean([pct_error(real_avg, dp_avg_capgain(0.1)) for _ in range(100)])
r2 = np.mean([pct_error(real_avg, dp_avg_capgain(1.0)) for _ in range(100)])
r3 = np.mean([pct_error(real_avg, dp_avg_capgain(10.0)) for _ in range(100)])

print("Average errors:", r1, r2, r3)

assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 20
assert r2 < 4
assert r3 < 0.4
```

Output:
```
Average errors: 3.008717582320781 0.29503100365680673 0.03440721335605419
```

## Question 4 (10 points)

In 2-5 sentences each, answer the following:

- Argue that your definition of `dp_avg_capgain` has a total privacy cost of `epsilon`
- For sums and averages, which seems to be more important for accuracy - the value of the clipping parameter $b$ or the scale of the noise added? Why?
- Do you think the answer to the previous point will be true for every dataset? Why or why not?

<!-- grade_id: cell-7ea16016366529bd, points: 10, grade: true, solution: true -->
- dp_avg_capgain must have a total privacy cost of epsilon * 2 in its current state, since I am using the differentially calculated sum from the dp_sum_capgain function and then also performing a laplace_mech on the averaged result.
- The scale of the noise added has to be more important since when extrapolated on a much larger data set, there is a much wider range of values to capture. Whereas the clipping parameter seemed to moreso help just bound the existing values.
- I do not think this point is consistent across all datasets because a database could contain data with a much wider and unpredictable range of values then those given in this assignment.

## Question 5 (20 points)

Write a function `auto_avg` that returns the differentially private average of a Pandas series `s`. Your function should automatically determine the clipping parameter `b`, and should enforce differential privacy for a **total privacy cost** of `epsilon`. You can assume that all values are non-negative (i.e. 0 or greater).

<!-- grade_id: cell-10b352b38a9b5f85, points: 0, grade: false, solution: true -->
```python
def auto_avg(s, epsilon):
    bs = list(range(min(s), max(s)))
    noisy_results = {}
    
    for b in bs:
        clipped_av = s.clip(lower=min(s), upper=b).mean()
        noisy_av = laplace_mech(clipped_av, b, epsilon)
        noisy_results[b] = noisy_sum
    
    last_av = -10000
    for b, noisy_av in noisy_results.items():
        if noisy_av < last_av:
            return b
        last_av = noisy_av
    return -1
```

<!-- grade_id: cell-49fb57b1ebfcc0da, points: 20, grade: true, solution: false -->
```python
# TEST CASE for question 5
r1 = np.mean([pct_error(adult['Age'].mean(), auto_avg(adult['Age'], 1.0)) for _ in range(20)])
r2 = np.mean([pct_error(adult['Capital Gain'].mean(), auto_avg(adult['Capital Gain'], 1.0)) for _ in range(20)])
r3 = np.mean([pct_error(adult['fnlwgt'].mean(), auto_avg(adult['fnlwgt'], 1.0)) for _ in range(20)])

print('Average errors:', r1, r2, r3)
assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 1
assert r2 < 100
assert r3 < 1
```

## Question 6

In 2-5 sentences each, answer the following:

- Explain your strategy for implementing `auto_avg`
- Argue informally that your definition of `auto_avg` has a total privacy cost of `epsilon`
- Did your solution work well for all three example columns? If it did not work well on any of them, why not?
- When is your solution likely to *not* work well? (i.e. what properties does the data have to have, in order for your solution to not work well?)

<!-- grade_id: cell-0ff63739ed66adf7, points: 10, grade: true, solution: true -->
- Firstly to implement auto_avg, I knew that the range of sensititivities is relative to the series' min and max bounds so I create a range of testing values with those arguments. Next for each sensitivity value in that range, an average clipped on the lower bound of the series and the respective upper bound is computed. Then noise is added to that result and saved to a dict for comparison. Those saved results are then compared to find the most ideal value (closest to zero).
- The total privacy cost for auto_avg depends on the differnce of the max and min value of the series s. this is because we compute a average using our privacy parmeter epsilon for every possible sensitivity value. Therefore the total privacy cost must be epsilon * (max(s) - min(s)).
- It appears my solution works well because it seems to be passing the assertion tests, although I think it might just be taking a long time to run, but it's late and I'm tired now so I have to throw in the debugging towel. 

```python

```

```python

```